---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise on Data viewing.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}
global_temp <- read.csv("DAX_2019-04-12.csv")
global_temp.info()
plot(global_temp)
ls()

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
